An initial report of c241,000 to 335,000 Year old Rock Engravings and their relation to Homo naledi in the Rising Star cave system, South Africa

Author response:

The following is the authors’ response to the original reviews

We thank the reviewers for their very constructive and helpful comments on the previous version of this manuscript. They have focused on some important issues and have raised many valuable questions that we expect to answer as research begins on these markings. As has been often the case with preprints, a number of experts beyond the four reviewers and editor have provided comments, questions, and suggestions, and we have taken these on board in our revision of the manuscript. In particular, Martinón-Torres et al. (2024) focused several comments upon this manuscript and raise some points that were not considered by the reviewers, and so we discuss those points here in addition to the reviewer comments.

Some of us have been engaged in other aspects of the possible cultural activities of Homo naledi. After the discovery of these markings we considered it indefensible to publish further research on the activity of H. naledi within this part of the cave system without making readers aware that the H. naledi skeletal remains occur in a spatial context near markings on cave walls. Of course, the presence of markings leaves many questions open. A spatial context does not answer all questions about the temporal context. The situation of the Dinaledi Subsystem does entail some constraints that would not apply to markings within a more open cave or rock wall, and we discuss those in the text.

We find ourselves in agreement with most of the reviewers on many points. As reflected by several of the reviewers, and most pointedly in the remarks by reviewer 1, the purpose of this preprint is a preliminary report on the observation of the markings in a very distinctive location. This initial report is an essential step to enable further research to move forward. That research requires careful planning due to the difficulty of working within the Dinaledi Subsystem where the markings are located. This pattern of initial publication followed by more detailed study is common with observations of rock art and other markings identified in South Africa and elsewhere. We appreciate that the reviewers have understood the role of this initial study in that process of research.

Because of this, the revised manuscript represents relatively minimal changes, and all those at the advice of reviewers. Many thanks to all the reviewers for noting various typographic errors, missed references and other issues that we have done our best to fix in the revised manuscript.

Expertise of authors. Reviewer 4 mentions that the expertise of the authors does not include previous publication history on the identification of rock art, and other reviewers briefly comment that experts in this area would enhance the description. AF does have several publications on ancient engravings and other markings; LRB has geological training and field experience with rock art. Notwithstanding this, we do take on board the advice to include a wider array of subject experts in this research, and this is already underway.

Image enhancement. We appreciate the suggestions of some reviewers for possible strategies to use software filters to bring out details that may not be obvious even with our cross-polarization lighting and filtering. These are great ideas to try. In this manuscript we thought that going very far into software editing or image enhancement might be perceived by some readers as excessive manipulation, particularly in an age of AI. In future work we will experiment with the suggested approaches.

Natural weathering. In the process of review and commentary by experts and the public there has been broad acceptance that many of the markings illustrated in this paper are artificial and not a product of natural weathering of the dolomite rock. We deeply appreciate this. At the same time, we accept the comments from reviewers that some markings may be difficult to differentiate from natural weathering, and that some natural features that were elaborated or altered may be among the markings we recognize. On pages 3 and 4 we present a description of the process of natural subaerial weathering of dolomite, which we have rooted in several references as well as our own observations of the natural weathering visible on dolomite cave walls in the Rising Star cave system. This includes other cave walls within the Dinaledi Subsystem. We discuss the “elephant skin” patterning of natural dolomite surface weathering, how that patterning emerges, and how that differs from the markings that are the subject of this manuscript.

Animal claw marks. Martinón-Torres et al. 2024 accept that some of the markings illustrated on Panel A are artificial, but they offer the hypothesis that some of those markings may be consistent with claw marks from carnivores or other mammals. They provide a photo of claw marks within a limestone cave in Europe to illustrate this point. On pages 5 and 6 of the revised manuscript we discuss the hypothesis of claw marks. We discuss the presence of animals in southern Africa that may dig in caves or mark surfaces. However the key aspect of the Malmani dolomite caves is that the hardness of dolomitic limestone rock is much greater than many of the limestone caves in other regions such as Europe and Australia, where claw marks have been noted in rock walls. As we discuss, we have not been able to find evidence of claw marks within the dolomite host bedrock of caves in this region, although carnivores, porcupines, and other animals dig into the soft sediments within and around caves. The form of the markings themselves also counter-indicates the hypothesis that they are claw marks.

Recent manufacture. One comment that occurs within the reviews and from other readers of the preprint is that recent human visitors to the cave, either in historic or recent prehistoric times, may have made these marks. We discuss this hypothesis on page 6 of the revised manuscript. The simple answer is that no evidence suggests that any human groups were in the Dinaledi Subsystem between the presence of H. naledi and the entry of explorers within the last 25 years. The list of all explorers and scientific visitors to have entered this portion of the cave system is presented in a table. We can attest that these people did not make the marks. More generally, such marks have not been known to be made by cavers in other contexts within southern Africa.

Panels B and C. We have limited the text related to these areas, other than indicating that we have observed them. The analysis of these areas and quantification of artificial lines does not match what we have done for the Panel A area and we leave these for future work.

Presence of modern humans. We have observed no evidence of modern humans or other hominin populations within the Dinaledi Subsystem, other than H. naledi. Several reviewers raise the question of whether the absence of evidence is evidence of absence of modern humans in this area. This is connected by two of the reviewers to the observation that the investigation of other caves in recent years has shown that markings or paintings were sometimes made by different groups over tens of thousands of years, in some cases including both Neanderthals and modern humans. We have decided it is best for us not to attempt to prove a negative. It is simple enough to say that there is no evidence for modern humans in this area, while there is abundant evidence of H. naledi there.

Association with H. naledi. Reviewer 2 made an incisive point that the previous version contained some text that appeared contradictory: on the one hand we argued that modern humans were not present in the subsystem due to the absence of evidence of them, yet we accepted that H. naledi may have been present for a longer time than currently established by geochronological methods.

We appreciate this comment because it helped us to think through the way to describe the context and spatial association of these markings and the skeletal remains, and how it may relate to their timeline. Other reviewers also raised similar questions, whether the context by itself demonstrates an association with H. naledi. We have revised the text, in particular on pages 5 and 7, to simply state that we accept as the most parsimonious alternative at present the hypothesis that the engravings were made by H. naledi, which is the only hominin known to be present in this space.

Age of H. naledi in the system. At one place in the previous manuscript we indicated that we cannot establish that H. naledi was only active in the cave system within the constraints of the maximum and minimum ages for the Dinaledi Subsystem skeletal remains (viz., 335 ka – 241 ka), because some localities with skeletal material are undated. We have adjusted this paragraph on page 7 to be clear that we are discussing this only to acknowledge uncertainty about the full range of H. naledi use of the cave system.

Geochronological methods. Several reviewers discuss the issue of geochronology as applied to these markings. This is an area of future investigation for us after the publication of this initial report. As some reviewers note, the prospects for successful placement of these engraved features and other markings with geochronological methods depends on factors that we cannot predict without very high-resolution investigation of the surfaces. We have included greater discussion of the challenges of geochronological placement of engravings on page 6, including more references to previous work on this topic. We also briefly note the ethical problems that may arise as we go further with potentially invasive, destructive or contact studies of these engravings, which must be carefully considered by not just us, but the entire academy.

Title. Some reviewers suggested that the title should be rephrased because this paper does not use chronological methods to derive date constraints for the markings. We have rephrased the title to reflect less certainty while hopefully retaining the clear hypothesis discussed in the paper.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: The production of painted, etched or engraved designs on cave walls or other surfaces is recognized as a major cognitive step in human evolution. Such intentional designs, which are widely interpreted as signifying, recording, and transmitting information in a durable manner were once considered exclusive to Late Pleistocene Homo sapiens. Recent work has demonstrated that other hominin groups also made such marks, including Neanderthals (Rodríguez-Vidal et al., 2014; Hoffmann et al., 2018), and possibly Middle-Pleistocene Homo erectus (Joordens et al., 2015). Such durable signs indicate an intentionality characteristic of meaning-making (Kissel and Fuentes 2018) which has been argued to require significant levels of cognitive abilities not found in species with smaller brain sizes (Parkington, 2010). In fact, the evolution of such meaning-making symbols is thought to be a core aspect of what it means to be “human” (Henshilwood, 2009). Here we present the first known example of abstract patterns and shapes engraved within the Dinaledi subsystem of the Rising Star Cave in South Africa. We identified markings incised into the dolomitic limestone walls of the cave. The engravings described here are deeply impressed cross-hatchings and other geometric shapes. The surfaces bearing these engravings appear to have been prepared and smoothed. In some areas there is residue that creates a sheen on the surface possibly indicating repeated handling or rubbing of the rock, and there is evidence of the application of dirt or sand to the surface by non-natural processes. Homo naledi entered this part of the cave system and buried bodies within the both the Dinaledi Chamber and adjacent Hill Antechamber between 241 and 335 ka (Dirks et al., 2017; Robbins et al., 2021, Berger et al, 2023a). The engravings described here are found on a pillar in the Hill Antechamber that extends into the natural fissure corridor that links the two chambers and we associate them with H. naledi.

Revised version: The production of painted, etched or engraved designs on cave walls or other surfaces is recognized as a major cognitive step in human evolution. Such intentional designs, which are widely interpreted as signifying, recording, and transmitting information in a durable manner were once considered exclusive to Late Pleistocene Homo sapiens. Here we present observations of what appear to be engraved abstract patterns and shapes within the Dinaledi Subsystem of the Rising Star cave system in South Africa, incised into the dolomitic limestone walls of the cave. The markings described here are found on a pillar in the Hill Antechamber that extends into the natural fissure corridor that links the two chambers and we associate them with H. naledi. They include deeply impressed lines, cross-hatchings, percussion marks, and other geometric shapes on flat wall surfaces and in and around existing cracks and grooves in the dolomitic limestone walls, found in one specific location of the Dinaledi Subsystem. Remains of multiple Homo naledi are found in this part of the cave system and evidence mortuary behaviour appears in both the Dinaledi Chamber and adjacent Hill Antechamber dated to between 241 and 335 ka (Dirks et al., 2017; Robbins et al., 2021, Berger et al, 2025).

Introduction

Earlier version: The Rising Star cave system, South Africa, is located within a small promontory situated to the south and east of the course of the Blaaubankspruit stream. The cave system is situated within the dolomitic limestone of the Malmani Subgroup, a Precambrian marine rock bedded with chert bands and containing abundant stromatolite fossils (Dirks et al. 2015; Eriksson et al. 2006). The system includes more than 3 km of mapped passages comprising multiple levels within a west-dipping dolomite horizon. Abundant remains of Homo naledi (Berger et al., 2015) occur within several localities in the system, including the Dinaledi subsystem, which lies at a depth of ∼30 m below the present surface and ∼120 m through the cave system from the nearest present entrance (Hawks et al, Elliot et al., 2019). Here, burials and other remains of H. naledi have been recovered and excavated from the Dinaledi Chamber, Hill Antechamber, and adjacent spaces and fissures (Berger et al. 2015; Berger et al. 2023; Brophy et al. 2021). These spaces are challenging to enter and navigate, and exploration of them is ongoing (Elliott et al. 2021).

Revised version: The Rising Star cave system, South Africa, is located within a small promontory situated to the south and east of the course of the Blaaubankspruit stream. The cave system is situated within the dolomitic limestone of the Malmani Subgroup, a Precambrian marine rock bedded with chert bands and containing abundant stromatolite fossils (Dirks et al. 2015; Eriksson et al. 2006). The system includes more than 3 km of mapped passages comprising multiple levels within a west-dipping dolomite horizon. Abundant remains of Homo naledi (Berger et al., 2015) occur within several localities in the system, including the Dinaledi Subsystem, which lies at a depth of ∼30 m below the present surface and ∼120 m through the cave system from the nearest present entrance (Hawks et al., 2017; Elliott et al., 2021). Remains of H. naledi have been recovered and excavated from the Dinaledi Chamber, Hill Antechamber, and adjacent spaces and fissures (Berger et al. 2015; Berger et al. 2025; Brophy et al. 2021). These spaces are challenging to enter and navigate, and exploration of them is ongoing (Berger et al. 2025, Elliott et al. 2021).

Earlier version: On July 28, 2022, during a survey of the Dinaledi Subsystem, we identified what appear to be engraved markings on the southern and northern faces of a natural pillar that forms the entrance and exit of a passage connecting the Hill Antechamber with the Dinaledi Chamber (Figure 1). Most of these marks are linear features between ∼5 and ∼15 cm in length. Many of these intersect to form geometric patterns such as squares, triangles, crosses, and X’s, while some are isolated lines. The engravings are located on three dolomitic panels, which we have labelled A, B and C. Seen as a triptych, these engravings are in a location where they can be viewed during access and egress to the Dinaledi Chamber when entering the system from the Hill Antechamber. The Hill Antechamber is the likely point of access by Homo naledi to the entire subsystem, and the passage is the natural linkage between the two main chambers of the subsystem (See Figure 1, also Elliott et al. 2021).

Revised version: On July 28, 2022, during a survey of the Dinaledi Subsystem, we identified what appear to be engraved markings on the southern and northern faces of a natural pillar that forms the entrance and exit of a narrow passage connecting the Hill Antechamber with the Dinaledi Chamber (Figure 1). The marks include linear features between ∼5 and ∼15 cm in length. Many intersect to form geometric patterns such as squares, triangles, crosses, and X’s, while some are isolated lines. These markings are located on three dolomitic panels on otherwise smooth areas, as well as in and around existing cracks and grooves in the rock. They occur in three areas which we have labelled A, B and C. These markings are in a location where they can be viewed during access into and egress from the Dinaledi Chamber when entering the system from the Hill Antechamber. The Hill Antechamber is the likely point of access by Homo naledi to the entire subsystem, and the passage is the natural linkage between the two main chambers of the subsystem (See Figure 1, also Elliott et al. 2021).

Added in the revised version: A photogrammetric map of the Dinaledi Subsystem of the Rising Star cave system, South Africa.

Earlier version: A photogrammetric map of the Dinaledi Subsystem of the Rising Star Cave system, South Africa. Orange bars mark the positions of the Engravings panels on the walls of the cave. Red boxes outline areas where excavations have been undertaken in the floor of the chambers. The green box outlines the area enlarged and colorized.

Revised version: Orange bars mark the positions of the Engravings panels on the walls of the cave. Red boxes outline areas where excavations have been undertaken in the floor of the chambers. The green box outlines the area enlarged and colorized.

Earlier version: In this paper, we describe detailed observations of Panel A within the passage linking the two main chambers. We present illustrations of Panels B and C within the Hill Antechamber and discuss their contextual relationship with Panel A, while recognizing that identifying all engraved lines within these panels will require further study in this difficult to access space. We also provide additional contextual data demonstrating the attribution of these etchings and engravings to H. naledi, hypothesise how the Panel A etchings and engravings were created, and discuss implications of our findings for H. naledi culture and cognition. We have not carried out any invasive or destructive sampling of these panels. This description is intended to document the discovery and provide spatial and contextual information prior to any further analyses that may require invasive sampling.

Revised version: In this paper, we describe our observations of Panel A within the passage linking the two main chambers. We provide images of Panels B and C within the Hill Antechamber and offer suggestions of their contextual relationship with Panel A, while recognizing that identifying and examining any all potential engraved lines within these panels requires further study. We also provide additional contextual data supporting the attribution of these markings to H. naledi, and discuss the time sequence represented by engraved marks on Panel A. We end by discussing potential implications of our findings for H. naledi culture and cognition. The description and initial analyses provided here is intended to document the discovery and provide spatial and contextual information to the scientific and heritage community prior to any further analyses that may require invasive sampling as such may require destructive methods or interacting with the marks directly, and therefore have associated ethical implications that need to be carefully considered before being undertaken (e.g. Bednarik, 2008; Arthur et al., 2021; Zerboni et al., 2022).

Panel A

Earlier version: Panel A is found on the southern face of the natural pillar that forms the southern edge of the entry from the Hill Antechamber into the southern of two passages leading to the Dinaledi Chamber (Figures 1, 2). The panel is notable as an area of discoloured rock that appears to have been smoothed by both percussive blows by a hard object, as is evidenced by micro and macro pitting of the surface alien to the adjacent natural rock surfaces (Figures 3 and 4) and by the possible application of sand and grit both before and after etchings and engravings were made (Figures 5 and 6). The adhering sediment and polishing of the surface of all three panels is unique to these surfaces relative to other surfaces in the chambers, and we thus hypothesize that it may result from intentional action. This sediment or pigment may have been used either as a material to create visual contrast on the grey dolomite, to abrade the surface as a form of polish, to enhance or obscure some aspects of the engraved lines, or all of these. This material is present on the surface as a micro-layer and is evident within some of the grooves of the lines, indicating its application after some of the marks were made. The appearance of time-ordering between engraved lines and the surface treatment may imply an origin of the engravings in multiple episodes (Figure 6).

Revised version: What we term Panel A is found on the southern face of the natural pillar that forms the southern edge of the entry from the Hill Antechamber into the southern of two passages leading to the Dinaledi Chamber (Figures 1, 2). The panel is notable as an area of discoloured rock that appears to have been smoothed by percussive blows by a hard object, as is evidenced by micro and macro pitting of the surface not found on the adjacent natural rock surfaces (Figures 3 and 4) and by the possible application of sand and grit both before and after markings were made (Figures 5 and 6).

Added in the revised version: Engraving Panel A (Images A and B) and Engraving Panel B (Images C and D).

Earlier version: Engraving Panel A (Images A and B) and Engraving Panel B (Images C and D). Image A is taken with a polarizing filter as described in methods. Image B taken using only LED lights and approximates natural coloration. Image C shows the results of increasing contrast while lowering light on Panel B while Image D illustrates Panel B under LED lighting.

Revised version: Image A is taken with a polarizing filter as described in methods. Image B taken using only LED lights and approximates natural coloration. Image C shows the results of increasing contrast while lowering light on Panel B while Image D illustrates Panel B under LED lighting.

Added in the revised version: Crosshatched etchings in Panel B.

Earlier version: Crosshatched etchings in Panel B. The white circle outlines areas of the engraving that may indicate hammer blows or pounding marks as evidenced by pitting not seen on other surfaces.

Revised version: The white circle outlines areas of the engraving that may indicate hammer blows or pounding marks as evidenced by pitting not seen on other surfaces.

Crosshatched etching comparing polarized images (bottom) with non-polarized imaging of the same area highlighting pitting marks that appear to be non-natural in origin.

Added in the revised version: Closeup of non-geometric figure at the top of Panel B.

Earlier version: Closeup of non-geometric figure at the top of Panel B. Note the cross like etching to the left of the figure as well as the X etched to the right. The non-geometric figure uses in part a natural fracture as an extension of the line beneath it before an inverted Y is etched at the terminus of this line. The material causing discoloration of the surface has not been analysed.

Revised version: Note the cross like etching to the left of the figure as well as the X etched to the right. The non-geometric figure uses in part a natural fracture as an extension of the line beneath it before an inverted Y is etched at the terminus of this line. The material causing discoloration of the surface has not been analysed.

Added in the revised version: Evidence for sediment covering lines on Panel A.

Earlier version: Evidence for sediment covering lines on Panel A. Image 1 shows lines 2,3 and 4 of Panel A (See Figure X for map). Note small sediment granules in base of lines to the left of the image, while line rising to the upper right shows penetration to the native underlying rock by the carving action. Image 2 illustrates a position slightly lower on the crosshatch marks on Panel A imaging lines 3,4,6,11 and 13. Note the difference between etching marks on the lower part and right of the image of this section of the engraving showing the difference between highly etched lines versus one presumably covered by a light layer of sediment post their creation. Note also likely pitting or presumed hammerstone marks in the central part of Image 2 between the carved lines. Image 3 illustrates lines 6, 17 and 18. Note the sharply carved lines on the right and the lines on the left that appear to be obscured by a light application of coarse sediment. Image 4 illustrates a wider shot of Panel B showing the discoloration of the area containing the engravings compared to the native rock with no sediment visible in the upper left of the image.

Revised version: Image 1 shows lines 2,3 and 4 of Panel A (See Figure X for map). Note small sediment granules in base of lines to the left of the image, while line rising to the upper right shows penetration to the native underlying rock by the carving action. Image 2 illustrates a position slightly lower on the crosshatch marks on Panel A imaging lines 3,4,6,11 and 13. Note the difference between etching marks on the lower part and right of the image of this section of the engraving showing the difference between highly etched lines versus one presumably covered by a light layer of sediment post their creation. Note also likely pitting or presumed hammerstone marks in the central part of Image 2 between the carved lines. Image 3 illustrates lines 6, 17 and 18. Note the sharply carved lines on the right and the lines on the left that appear to be obscured by a light application of coarse sediment. Image 4 illustrates a wider shot of Panel B showing the discoloration of the area containing the engravings compared to the native rock with no sediment visible in the upper left of the image.

Added in the revised version: Boc like etchings on Panel B seen under ultraviolet light.

Earlier version: Boc like etchings on Panel B seen under ultraviolet light. Note the slight white appearance of the etched lines indicating the presence of a reflective or slightly fluorescing material in the engraved lines similar to the fluorescence of pure CaCO3.

Revised version: Note the slight white appearance of the etched lines indicating the presence of a reflective or slightly fluorescing material in the engraved lines similar to the fluorescence of pure CaCO3.

The Dinaledi Subsystem etched crosshatch found on Panel A between the Hill Antechamber and Dinaledi Chamber (Top) compared to the crosshatch engravings found on the cave floor of Gorham’s Cave, Gibraltar and attributed to manufacture by a neanderthal circa 60k years ago.

Earlier version: Line tracing of the Panel A Dinaledi Subsystem engraving (top in grey), compared to a line tracing of the Gorham’s Cave engraving (bottom in black). Tracings not to scale.

Revised version: Line tracing of the Panel A Dinaledi Subsystem engraving (top in grey), compared to a line tracing of the Gorham’s Cave engraving (bottom in black).

Added in the revised version: Tracings not to scale.

Earlier version: The most visible engraved lines, when viewed together are crosshatched, give the impression of a rough hashtag figure (Figures 10 and 11). The lines appear to have been made by repeatedly and carefully passing a pointed or sharp lithic fragment or tool into the grooves. This excludes the possibility of an unintentional or utilitarian origin. In addition, there are scratches that fall outside of identifiable designs, which may either be mistakes, unfinished designs, or form part of the design not interpretable by us. Several of the grooves overlap geological features native to the rock including fossil stromatolites (Figures 12 and 13). In many instances, it is possible to identify which lines were made first by examining the point where they cross another line (e.g. Figures 14 and 15). As has been interpreted for other discoveries of early geometric shapes etched or engraved by larger-brained species, this discovery demonstrates the capacity of H. naledi for expression through the use of geometric forms.

Revised version: The most visible markings are engraved lines, some of which when viewed together are crosshatched and give the impression of a rough hashtag figure (Figures 10 and 11). The lines appear to have been made by repeatedly and carefully passing a pointed or sharp incising tool into existing or new grooves. In addition, there are markings that are best described as scratches that fall outside of identifiable lines and their intersections. Several of the grooves overlap geological features native to the rock including fossil stromatolites (Figures 12 and 13). In many instances, we suggest it is possible to identify which lines were made first by examining the point where they cross another line (e.g. Figures 14 and 15).

Earlier version: Non-polarized (left) and polarized (right) image of the crosshatched engraving on Panel A, Dinaledi Subsystem. Scale in millimeters.

Revised version: Non-polarized (left) and polarized (right) image of the crosshatched engraving on Panel A, Dinaledi Subsystem.

Added in the revised version: Scale in millimeters.

Earlier version: Polarized image of the Panel A engraving (left) with the most visible lines (interpreted as the most recent etchings) traced (right). Scale in millimeters.

Revised version: Polarized image of the Panel A engraving (left) with the most visible lines (interpreted as the most recent etchings) traced (right).

Added in the revised version: Scale in millimeters.

Earlier version: LED light images of the Panel A primary engravings. Scales in millimetres.

Revised version: LED light images of the Panel A primary engravings.

Added in the revised version: Scales in millimetres.

Added in the revised version: Lines 2,6,11, 14, 15,16,18, 21,25,27,33,38,39 and 40.

Earlier version: Lines 2,6,11, 14, 15,16,18, 21,25,27,33,38,39 and 40. Note that etching 18 overcuts line 17 and line 30 overcuts line 25 indicating an order of creation. Note also that all line engraved on the left side of the image cut through the fossil stromatolite visible as horizontal wavy lines in the rock.

Revised version: Note that etching 18 overcuts line 17 and line 30 overcuts line 25 indicating an order of creation. Note also that all line engraved on the left side of the image cut through the fossil stromatolite visible as horizontal wavy lines in the rock.

Added in the revised version: Magnified views of Lines 6, 17 and 18.

Earlier version: Magnified views of Lines 6, 17 and 18. The bottom image is a slightly higher magnification of the top image. It is clear from these images that the engraving of line 17 preceded the engraving of line 18. Scale at bottom in millimetres.

Revised version: The bottom image is a slightly higher magnification of the top image. It is clear from these images that the engraving of line 17 preceded the engraving of line 18. Scale at bottom in millimetres.

Added in the revised version: Magnified views of lines 16, 25 and 30.

Earlier version: Magnified views of lines 16, 25 and 30. Note it is clear line 25 was etched first, followed by line 30. The lateral edge of line 31 can be seen at bottom left and the left edge of line 15 at top right. Note also the deep incision through the stromatolites layers by all lines. Scale in millimeters.

Revised version: Note it is clear line 25 was etched first, followed by line 30. The lateral edge of line 31 can be seen at bottom left and the left edge of line 15 at top right. Note also the deep incision through the stromatolites layers by all lines. Scale in millimeters.

Earlier version: We identify at least 46 non-natural engraved marks on panel A (See Figure 16). The most prominent markings on Panel A are a series of intersecting lines (Figures 10,11 and 12). There appears to be a temporal span involved in the creation of the engraved lines as some seem more recently engraved and show clean etching, while others have been obscured either by slight weathering or by the application of sediment. The most easily identifiable engravings, based on their clarity, are Lines L2, L6, L9, L11, L16, L17, L27, L30 and L31 (Figure 16). While the existing lines may have been created in older etchings, or been created over multiple interactions, the final etchings of the lines based on which lines overlap can be interpreted as follows: horizontal Lines 11 & 25 were created after vertical Line 2. Vertical Line 6 was created after L11. Vertical Line 18 was created after horizontal Line 17. Line 30 was created after horizontal Line 25, but before horizontal Line 31.

Revised version: We identified at least 46 engraved or altered marks we consider artificial on panel A (See Figure 16). In this number we include features that appear to be natural lines that have been altered by etching. The most prominent markings on Panel A are a series of intersecting lines (Figures 10,11 and 12). There appears to be a temporal span involved in the creation of the engraved lines as some seem more recently engraved and show clean etching, while others have been obscured either by slight weathering or by the application of sediment. The most easily identifiable engravings, based on their clarity, are Lines L2, L6, L9, L11, L16, L17, L27, L30 and L31 (Figure 16). While the existing lines may have been created in older etchings, or been created over multiple interactions, the final etchings of the lines based on which lines overlap we interpret as follows: horizontal Lines 11 & 25 were created after vertical Line 2. Vertical Line 6 was created after L11. Vertical Line 18 was created after horizontal Line 17. Line 30 was created after horizontal Line 25, but before horizontal Line 31.

Earlier version: A conservative map of non-natural engravings observed on Panel A. Non-natural engravings are traced in white lines and given yellow numbers references in the text.

Revised version: A conservative map of non-natural engravings observed on Panel A.

Added in the revised version: Non-natural engravings are traced in white lines and given yellow numbers references in the text.

Evidence of hominin manufacture of engravings on Panel A

Earlier version: Dolomite is known for a pattern of natural weathering that results in patterns of recessed linear features on its surface. Artificial markings can be distinguished from this natural weathering pattern in several ways. Natural fissures and erosional features in weathered dolomite surfaces are characteristically deeper than several millimeters and they follow natural fracture planes within the rock. Artificial lines are limited in depth and extent due to the natural hardness of dolomite. This hardness means that any substantial artificial marking requires multiple parallel incisions with a hard tool. Natural erosional features in dolomite may have variable cross-sections, ranging from bevelled to U-shaped to rectangular in cross section, but do not have multiple parallel striations visible within them. Where artificial engraved markings intersect, they often exhibit an ordering in which one was completed before the other; this kind of feature is not typical of natural weathering. In previous work, researchers have noted the limited depth of artificial lines, their composition from multiple parallel striations, and their association into a clear arrangement or pattern as evidence of hominin manufacture (Fernandez-Jalvo et al. 2014).

Revised version: Dolomite is known for a pattern of natural weathering that results in patterns of recessed linear features on its surface often called “elephant skin” patterning. Natural fissures and erosional features in weathered dolomite surfaces due to chemical weathering processes (Burger, 1989) are in our observations characteristically deeper than several millimeters and they follow natural fracture planes within the rock. Natural erosional features in dolomite may have variable cross-sections, ranging from bevelled to U-shaped to rectangular in cross section with rounded or flat bottoms (Shah, Jeelani and Goldscheider, 2018) and are typically a few millimeters to many centimetres in depth, but do not have multiple parallel striations visible within them which would be characteristic of retouching or engraving. As opposed to natural surface weathering, artificial lines are limited in depth and extent due to the natural hardness of dolomite. This hardness (see Tabor, 1954) means that any substantial artificial marking requires multiple parallel incisions with a hard tool. Where artificial engraved markings intersect, they often appear to exhibit an ordering in which one was completed before the other (see for example figures 6, 13, 14 and 15); this appearance of sequencing is not typical of natural weathering. In previous work, researchers have noted the limited depth of artificial lines, their composition from multiple parallel striations, and their association into a clear arrangement or pattern as evidence of hominin manufacture (Rodríguez-Vidal et al. 2014).

Earlier version: The engraved lines in Panel A have each of these features. They can clearly be distinguished from natural weathering of the surrounding dolomite walls, which can be seen adjacent to the panel within 20 cm of the nearest artificial marks (Figure 18). The features produced by natural weathering are deeper than 10 mm, in particular deep relative to the feature width, they maintain a consistency of size and depth across substantially undulating or rugged surfaces, they expand from natural cracks and fissures. In contrast, even the widest of the engraved lines that constitute Panel A have a relatively shallow depth. High-resolution macro-photography shows micro-striations constituting several of these engraved lines, in which roughly parallel incisions sometimes overlap with each other (See Figures 6, 13, 14, 15 and 17). Many of the lines also fall out of the direction of natural fracture features in the country rock, although it should be recognised that there are multiple places on this panel where natural lines and features of the rock may have been enhanced by artificial engraving. Figures 13, 14, and 15 show examples of ordering where engraved lines intersect, one having been completed clearly before the other.

Revised version: The engraved lines in Panel A diverge in their patterning from the natural “elephant skin” weathering of the surrounding dolomite walls, which can be seen adjacent to the panel within 20 cm of the nearest artificial marks (Figure 18). Our assessment of the features produced by natural weathering in this area is that they are deeper than 10 mm, in particular deep relative to the feature width, they maintain a consistency of size and depth across substantially undulating or rugged surfaces, they expand from natural cracks and fissures. In contrast, even the widest of the engraved lines that constitute Panel A have a relatively shallow depth. High-resolution macro-photography shows micro-striations constituting several of these engraved lines, in which roughly parallel incisions sometimes overlap with each other (See Figures 6, 13, 14, 15 and 17). Many of the lines also fall out of the direction of natural fracture features in the dolomite, although it should be recognised that there are multiple places on this panel where natural lines and features of the rock may have been enhanced or extended by artificial engraving. Figures 13, 14, and 15 show examples of ordering where engraved lines intersect, one having been completed clearly before the other. Some of the elements we describe as engraved lines appear to be artificial extensions or alterations of lines that may originally have resulted from natural weathering.

Earlier version: Magnified images of etchings 41 through 46 numbered from left to right. Scale in millimeters.

Revised version: Magnified images of etchings 41 through 46 numbered from left to right.

Added in the revised version: Scale in millimeters.

Added in the revised version: Image of dolomite above and right of Panel A.

Earlier version: Image of dolomite above and right of Panel A. The top of the crosshatched etchings can be seen in the lower left of the image. Note the smoothing and alteration of the Panel’s surface compared to natural, non-altered dolomitic surfaces above and right of the Panel typical of unaltered surfaces throughout the system.

Revised version: The top of the crosshatched etchings can be seen in the lower left of the image. Note the smoothing and alteration of the Panel’s surface compared to natural, non-altered dolomitic surfaces above and right of the Panel typical of unaltered surfaces throughout the system.

Earlier version: In addition to the engraving depth, composition, and ordering, there are two additional aspects of Panel A engraved lines that distinguish them from natural weathering. The dolomitic bedrock of the Malmani Formation includes fossil stromatolites, which manifest as curving linear banded striations visible in the rock. Panel A includes these layered stromatolitic bands, and all engraved lines that pass below the bottom of Line 14 cross over this fossil feature (See Figures 13 and 15). Where engraved lines cross over this feature, they retain direction and in some cases the multiple striations slightly diverge, suggesting that maintaining a linear engraving over this irregular surface may have been challenging. Second, the engraved markings are, in places, covered wholly or partially in sediment or some other substance. This coating on the walls of the cave does not occur in other areas of the chambers where there are no engravings. Thus it does not appear this covering sediment can be explained by geological or other non-organic processes.

Revised version: In addition to the engraving depth, composition, and ordering, two additional aspects of Panel A markings distinguish them from natural weathering. The dolomitic bedrock of the Malmani Formation includes fossil stromatolites, which manifest as curving linear banded striations visible in the rock. Panel A includes these layered stromatolitic bands, and all engraved lines that pass below the bottom of Line 14 cross over this fossil feature (See Figure s13 and 15). Where engraved lines cross over this feature, they retain direction and in some cases the multiple striations slightly diverge, suggesting that maintaining a linear engraving over this irregular surface may have been challenging. Second, the markings are, in places, covered wholly or partially in sediment or some other substance. This coating does not occur in other areas of the chambers where there are no such engravings. This material (we cannot determine if it is sediment or pigment) may have been used either to create visual contrast on the grey dolomite, to abrade the surface as a form of polish, to enhance or obscure some aspects of the engraved lines, or all of these. This material is present on the surface as a thin layer and is evident within some of the grooves of the lines, indicating its application after some of the marks were made. The appearance of time-ordering between engraved lines and the surface treatment may imply an origin of the engravings in multiple episodes (Figure 6).Thus it does not appear this covering, likely sediment, can be explained by geological or other non-organic processes.

Earlier version: The means of manufacture of these engraved lines would have required an implement of equal or greater hardness as the native dolomitic limestone. At present, only one possible lithic artifact has been recovered in direct association with H. naledi remains (Berger et al. 2023a). This tool-shaped rock does resemble tools from other contexts of more recent age in southern Africa, such as a silcrete tool with abstract ochre designs on it that was recovered from Blombos Cave (Henshilwood et al. 2018) (Figure 19). Dolomite rocks of appropriate size and morphology to mark the cave walls have been recovered from surface contexts within the Dinaledi Subsystem, as have many chert fragments.

Revised version: The means of manufacture of engraved lines in this context would have required an implement of equal or greater hardness compared to the native dolomitic limestone. Dolomite rocks of appropriate size and morphology to mark the cave walls have been recovered from surface contexts within the Dinaledi Subsystem, as have many chert fragments. At present, only one possible lithic artifact has been described in direct association with H. naledi remains (Berger et al. 2025). This “tool-shaped rock” does resemble tools from other contexts of more recent age in southern Africa, such as a silcrete artifact with linear engraved features on it that was recovered from Blombos Cave (Henshilwood et al. 2018) (Figure 19).

The tool-shaped artefact described in Berger et al, 2023a (top) recovered from the Hill Antechamber burial immediately below Panels B and C compared to the artefact from Blombos cave, South Africa attributed by Henshilwood et al 2009 as having symbolic markings in ochre made by Homo sapiens circa 78k years ago.

Panels B and C

Earlier version: Panels B and C are located on the northern wall within two meters of the Hill Antechamber burial feature described in Berger et al. (2023). Panel B is situated lower and to the right (West) of Panel C. Both panels appear to have been prepared in a similar way to Panel A, with possible use of cave sediment applied to the surface, giving the surfaces of these panels an obvious textural difference to adjacent walls of the chambers (Figures 2c & b and Figure 5) A number of obvious etchings and engravings can be seen, some in the form of geometric figures, crosses, X’s and one possible non-linear geometric figure (Figure 5). It appears, in softer visible light, that a foreign substance has been applied to part of the panel. As was noted the purpose of this paper is not to describe these complex panels and the many etchings and engravings on them, but to simply note their presence in the Hill Antechamber. Future work in this difficult space is planned to sample the possible residues and map the non-natural etchings, attempt to date the etchings and we will conduct experimental work on native dolomite in controlled experiments.

Revised version: Panels B and C are located on the northern wall within two meters of the Hill Antechamber burial feature described in Berger et al. (2023). Panel B is situated lower and to the right (West) of Panel C. Both panels appear to have been prepared in a similar way to Panel A, with possible use of cave sediment applied to the surface, giving the surfaces of these panels an obvious textural difference to adjacent walls of the chambers (Figures 2c & b and Figure 5) A number of etchings and engravings can be seen, some in the form of geometric figures, crosses, X’s and one possible non-linear geometric figure (Figure 5). It appears, in softer visible light, that a foreign substance has been applied to part of the panel. As was noted the purpose of this paper is not to describe these complex panels, but to simply note their presence in the Hill Antechamber. Future work in this difficult space is planned to sample the possible residues and map the non-natural etchings, attempt to date the etchings and we will conduct experimental work on native dolomite in controlled experiments.

Added in the revised version: Discussion

Added in the revised version: The context of these markings within the remote Dinaledi Subsystem, far from cave entrances and challenging to access (Dirks et al, 2015; Kruger et al., 2016; Elliott et al; 2021), suggests that the most parsimonious explanation for their manufacture is that Homo naledi produced them (Berger et al., 2015; Berger et al., 2024). This is the only species that left any evidence of presence in this subsystem until the entry of modern cave explorers. In this discussion we consider possible objections to this hypothesis and ways in which this hypothesis may be tested in the future. We also suggest some comparisons with the broader record of ancient markings beyond the Rising Star examples.

Added in the revised version: One possible objection is that some markings were not produced by hominins. For example, some have suggested that some of the apparent markings are consistent with weathered dolomite, while others may have been produced by non-hominin agents such as carnivores (Martinón-Torres et al., 2024). Linear features and other modifications of rock surfaces are sometimes produced by animal claws due to tunnelling, climbing, or claw-sharpening, or other behaviors (Bednarik, 1991). Such features have often been observed within caves formed within relatively soft limestone formations, including those in some parts of Europe and southern Australia.

Added in the revised version: Deeply incised claw markings are sometimes seen, which tend to include multiple parallel incisions with a spacing typical of the anatomical spacing between claws. When claw marks are made repeatedly on the same surface, they may cross each other or exhibit complex geometries. However, the susceptibility of a surface to marking by animal claws is determined by the hardness of the material. Animal claws have a hardness of 2 to 2.5 on Mohs hardness scale (Rothschild et al., 2013). Dolomite has a hardness of 3.5 to 4.5, with Malmani dolomites at the higher end of this scale, much harder than non-dolomitic limestones with hardness of 2 to 3 which are commonly marked by animal claws. None of the cave systems in the Malmani dolomite have been noted for incised claw markings into dolomite itself of the depth and form seen in the Dinaledi marks (Figure 17), even though clawed species like porcupines, honey badgers, and leopards do use caves and do dig into softer sediments found in the area. The hardness of the rock does not admit scratches by animal claws of this kind; additionally the specific features of Panel A do not exhibit the kind of spacing or morphology that would be expected for repeated scratching by animal claws, and their position well above the cave floor would not lend itself naturally to scratching in that location. Further, there is no evidence of carnivore presence within the subsystem: no carnivore remains, no evidence of carnivore burrowing activity, no coprolites, and no bones with tooth marks attributable to carnivores (Dirks et al., 2015; Dirks et al., 2016).

Added in the revised version: Another possible objection is that the engravings may be of comparatively recent origin, potentially made by modern cavers or other H. sapiens groups in Holocene or Late Pleistocene times (Herries, 2023; Petraglia et al., 2023; Martinón-Torres et al., 2024; Pettitt and Wood, 2024). This idea has two parts: one concerns the idea that other species or populations besides Homo naledi were active within this subsystem; the other concerns chronological placement of the markings relative to Homo naledi. There is no evidence of any hominin activity in the Dinaledi Subsystem other than H. naledi. No remains attributable to other prehistoric hominin groups occur in this subsystem or any adjacent areas of the cave system, nor is there any other evidence of activity of such groups (Dirks et al., 2015; Dirks et al., 2016). Human cave explorers entered these areas only within the last 40 years, and the number of modern cavers and archaeologists who have entered the Dinaledi Subsystem is extremely limited (Table 1). There is no evidence of modern cavers altering cave walls in such a manner in the Dinaledi Subsystem, or elsewhere in Rising Star system, and we have detected no evidence of traces of metallic tools in association with the markings. By contrast abundant remains of H. naledi have been identified within the subsystem including the Hill Antechamber Feature only three to four meters away (Berger et al., 2024). Nevertheless, the nature and composition of the tools that made these engravings are an area that can be tested with future work.

Added in the revised version: It is very challenging to establish the chronological age of engraved markings on rock surfaces (Van Riet Lowe, 1968; Butzer et al., 1979; Bednarik, 2002,2008; Whitley, 2012). Successful use of geochronological methods may occur when some or all of the markings were buried within sediment that can be dated (Rodriguez-Vidal et al. 2014), or when markings were later covered in part by mineral formations that can be dated using radiometric methods (Hoffmann et al., 2016), or when exposure to weathering creates a varnish that can be dated (Dom and Nobbs, 1992; Dragovich, 2000; Mason, 2006; Pillans and Faifield, 2013). Still, even where geochronological methods are possible, they may provide only a minimum age, which may be thousands of years after the marks were produced. The markings that we have identified within the Dinaledi Subsystem are not overlain by naturally forming sedimentary deposition. Without further excavation of sediments that adjoin cave walls we cannot say whether additional markings may exist in subsurface contexts. No obvious mineral deposition overlaps the engraved features. Our observations do note that there may be materials added to the surface of the engraved areas, and these may provide an avenue of investigation. The situation as a whole makes it challenging to build alternative lines of evidence about the age of the engravings beyond their spatial proximity to the Homo naledi evidence from only a few meters away (Berger et al. 2023).

Removed in the revised version: Discussion and Conclusions

Removed in the revised version: The attribution of engraved or painted markings to Neandertals, Homo erectus, or other hominin groups has generally attracted debate. Critics have emphasized the need to establish clearly the intentionality of possible markings in contrast to natural processes. Skepticism has also frequently surrounded methods to establish the geological age of engraved or painted markings (e.g., Pons-Branchu et al. 2020; White et al. 2020). Some have emphasized that while singular occurrences may indicate intentionality, only repeated evidence from multiple sites can provide evidence of possible symbolic or representational intent (Davidson 2020).

Earlier version: Geochronological evidence can be extremely difficult to obtain for markings on natural rock surfaces. The engraved panels in the Dinaledi subsystem are not overlain by sediments, and we have not identified any calcite formation overlapping the engraved features. This makes it challenging to assess whether the engravings are contemporary with the Homo naledi burial evidence from only a few meters away (Berger et al. 2023). At present we have no evidence limiting the time period across which H. naledi was active in the cave system. The maximum age constraint reported by Dirks et al. (2017) on H. naledi skeletal material (335 kyr BP) in Dinaledi is the highest 95% confidence limit of a direct ESR-US date on H. naledi teeth; while the minimum age constraint (241 kyr BP) is based on U-Th on a flowstone that formed in part around a bone fragment (Wiersma et al. 2020). These dates do not necessarily pertain to skeletal material from other parts of the cave system, nor do they exclude earlier or later access to the cave system by H. naledi individuals. The duration of H. naledi cultural activity within the cave system is therefore not presently known.

Revised version: Thus, at present the H. naledi skeletal material provides the only information about the chronology of hominin activity in this area of the cave system. The maximum age constraint reported by Dirks et al. (2017) on H. naledi skeletal material (335 kyr BP) in Dinaledi is the highest 95% confidence limit of a direct ESR-US date on H. naledi teeth; while the minimum age constraint (241 kyr BP) is based on U-Th on a flowstone that formed in part around a bone fragment (Dirks et al., 2017; Robbins et al., 2021). These dates provide strong constraints on some H. naledi skeletal material from the Dinaledi Subsystem. Geochronological methods do not yet provide similar constraints for H. naledi material from other parts of the cave system, including the Lesedi Chamber (Hawks et al. 2017). Therefore, we cannot be certain that the constraints of the Dinaledi Subsystem evidence mark the full duration of H. naledi cultural activity across the cave system. However, we adopt as a conservative hypothesis that the age of these markings is constrained in the same way as the skeletal remains from the Dinaledi Subsystem, between 335 ka and 241 ka.

Added in the revised version: Stylistic criteria are also sometimes used to establish the chronology of cave art, although such comparisons should be accepted cautiously. In this case, the images are not typical of contemporary graffiti made by recent people. They are also not typical of rock art of the last several thousand years in southern Africa. The markings are comparable to other hominin-produced markings, pigment use, and curation of nonutilitarian objects from the later Middle Pleistocene of southern Africa. Small slabs of ironstone with incised lines were excavated from Wonderwerk Cave, with a minimum age of 276 kyr BP (Beaumont and Vogel 2006; Beaumont and Bednarik 2013; Chazan and Horwitz 2009; Jacobson et al. 2012). One of these slabs bears a series of seven roughly parallel engraved lines, while crossed lines occur on other slabs (Beaumont and Bednarik 2013; Chazan and Horwitz 2009). Beaumont and Bednarik (2015) proposed that two sites in the southeastern Kalahari with open air rock engravings, Potholes Hoek and Nchwaneng, have highly eroded cupules dating to the early Middle Stone Age, possibly as early as 400 kyr. Small pebbles of quartz, chalcedony, and other materials occur within the Wonderwerk sequence and have been interpreted as manuports (Beaumont and Vogel 2006; Beaumont and Bednarik 2013), although an alternative hypothesis that they were ostrich gastroliths, potentially introduced by hominins who hunted the birds, may also be possible (Tryon 2010). Mineral pigments including specularite and hematite have been found in a number of contexts dating to earlier than 300 kyr, including Kathu Pan 1, Wonderwerk, and Canteen Kopje (Watts et al. 2016). None of these examples of hominin-produced markings, pigment use, or other activities have been found near hominin fossil remains. Some of them predate any evidence of modern humans or early Homo sapiens anywhere on the continent.

Removed in the revised version: It is unlikely that any other hominin population made these engravings. No physical or cultural evidence of any other hominin population occurs within this part of the cave system, and there is no evidence that recent humans or earlier hominins ever entered any adjacent area of the cave until surveys by human cave explorers during the last 40 years. The number of modern cavers and archaeologists who have entered the Dinaledi subsystem is extremely limited (Table 1). There is no evidence of modern cavers altering cave walls in such a manner in the Dinaledi subsystem, or elsewhere in Rising Star system. The evidence that these engravings were created in multiple events over time further makes it unlikely that historic humans were involved in their creation. The available evidence is most compatible with the extinct species Homo naledi as the creator of these markings.

Removed in the revised version: The evidence of burials and associated mortuary practices by H. naledi near the engravings reinforces that assertion this species carried out repeated complex patterns of behaviour in this deep cave setting (Berger et al. 2023a, Fuentes et al. 2023). The engravings are located in a distinctive position, on the left-hand wall as seen when entering the system from the North, and interior left hand pillar that forms the entrance archway to the tunnel linking the Hill Antechamber burial area with the larger Dinaledi Chamber burial area. This is the only place engravings have been discovered so far within the Dinaledi subsystem. The evidence that Panel A was marked in multiple episodes, possibly separated by substantial time, suggests that the selection of this location was not random, and that an individual or individuals returned to this location to carry out a similar pattern of activity on multiple occasions.

Earlier version: The main engravings on Panel A appear similar to other engravings found in the later Pleistocene. The shapes of the engravings on panels A,B and C also appear to include the following geometric forms identified by Von Petzinger (2017): crosshatch, cruciform, line, flabellifrom, scalariform, open angle and oval. However, further analytic and comparative work must be conducted to confirm exactly how much similarity and overlap there is between the Dinaledi engravings and the engravings at other Pleistocene sites where such designs are found. As a specific observation however, the engravings in panel A give the impression of overlapping crosses and lines and are remarkably similar in appearance to the engraving from Gorham’s Cave, Gibraltar (Rodriguez-Vidal et al. 2014). This engraving was dated to greater than 39 kyr cal BP and has been attributed to Neandertals. Other geometric patterns made with lines occur in several contexts are reported for some later Pleistocene sites in southern Africa and elsewhere (e.g. Von Petzinger, 2017). These include ochre lines, engraved bones, and engraved ochre chunks from Blombos Cave (d’Errico et al. 2001; Henshilwood et al. 2002; Henshilwood et al. 2018), engravings from Wonderwerk Cave (Thackeray et al. 1981) and lines impressed within sand features that were later lithified into aeolianites (Helm et al. 2021). There are also a few other engravings from sites in Europe at similar time depth (Von Petzinger, 2017; Kissel and Fuentes 2017, 2018), as well as geometric lines on a freshwater mussel shell from Trinil Java, attributed geochronologically to H. erectus (Joordens et al. 2015). The engravings from the Dinaledi Subsystem share similarities with many of these geometric expressions from other sites and geographic regions. The Blombos artifacts also include some surfaces that appear to have been prepared or smoothed prior to engraving possibly similar to the processes involved in the smoothing of Dinaledi Panel A.

Revised version: The Dinaledi markings are comparable to some of the incised lines and engravings found at other sites in the later Middle and Late Pleistocene. Panel A gives the impression of overlapping crosses and lines that are remarkably similar in appearance to the engraving from Gorham’s Cave, Gibraltar (Rodriguez-Vidal et al. 2014). This engraving was dated to greater than 39 kyr cal BP and has been attributed to Neandertals. In addition to the incised slabs and cupules mentioned above, the later Pleistocene record of southern Africa includes engraved bones and engraved ochre blocks from Blombos Cave (d’Errico et al. 2001; Henshilwood et al. 2002; Henshilwood et al. 2018) and lines impressed within sand features that were later lithified into aeolianites (Helm et al. 2021). There are also a few other engravings from sites in Europe at similar time depth (Von Petzinger, 2017; Kissel and Fuentes 2017, 2018), as well as geometric lines on a freshwater mussel shell from Trinil, Java, attributed to H. erectus based on geochronological considerations (Joordens et al. 2015). The markings in the Dinaledi Subsystem share similarities with many of these geometric expressions from other sites and geographic regions. Several geometric forms identified by Von Petzinger (2017) are present, including crosshatch, cruciform, line, flabellifrom, scalariform, open angle and oval. However, further analytic and comparative work must be conducted to confirm exactly how much similarity and overlap there is between the proposed Dinaledi engravings and the engravings at other Pleistocene sites where such designs have been found.

Earlier version: Many of these examples of engraved lines from later Pleistocene sites appear to be nonrandomly placed on an object or surface. Henshilwood and Dubreuil (2011) have suggested that one should be less focused on the specifics of the designs and rather concentrate on the underlying cause of their creation. Those and other authors suggest symbolic implications for such engravings and associated them with the emergence of contemporary Homo sapiens. However, the recent identification of engravings and other forms of material meaning making in a range of other-than-Homo sapiens hominins over the latter portions of the Pleistocene (Kissel and Fuentes 2018, 2021) suggests that such activity, be it “symbolic” or not, is not exclusive to Homo sapiens. With the engravings reported here we add to this growing dataset by providing additional evidence of later Pleistocene engravings associated with a non-Homo sapiens hominin. We also add to the complexity involved in examining and understanding the implications of such engravings by reporting that the most likely creator of these engravings was the small-brained Homo naledi. This has implications for the evolution of biological intelligence among hominins and the association with encephalization with cognitive complexity.

Revised version: The precise location of the engraved markings within the Dinaledi Subsystem may be important to their interpretation. Many artifacts and rock walls from later Pleistocene sites with engraved marks appear to have a nonrandom placement of markings on an object or surface. Henshilwood and Dubreuil (2011) have suggested that one should be less focused on the specifics of the designs and rather concentrate on the underlying cause of their creation. Those and other authors have suggested symbolic implications for such engravings, under the assumption that they were made by Homo sapiens. The broader record of engravings and other forms of meaning-making behavior in a range of hominin species and populations over the latter portions of the Pleistocene (Kissel and Fuentes 2018, 2021) suggests that such activity, be it “symbolic” or not, was not exclusive to Homo sapiens. The engravings reported here add to this growing dataset.

Methods

Earlier version: The etchings and engraving markings were examined using high resolution photography and magnification of lines and markings. Polarizing filters were also used to enhance relief and this is indicated when used.

Revised version: As the main purpose of this work was documentation of the presence of engravings in proximity to the burials, for this study, the etchings and engraving markings were only examined using high resolution photography and magnification of lines and markings. Polarizing filters were also used to enhance relief and this is indicated when used.

Cross-polarisation was employed for control of specular highlights/reflections in order to limit artefacts when generating the 3D-depth map for photogrammetry purposes. A circular polariser was used on the camera lens in conjunction with a linear polarising gel placed over the two speed lights (electronic flash heads) used as the light source. The different minerals/material on the dolomite are reflecting/absorbing the cross-polarised light emphasising the “bright” striations visible in images.

Images were shot with a 50mm (Polariser fitted) at f/11 unless otherwise stated.

The light source used (twin speed lights with polarised gel attached) were placed as close to the Lens axis as possible so that the angles of incidence approximate the reflected angles limiting shadow. This assisted us in building the 3D mesh for photogrammetry purposes. The cross-polarisation also removed specular highlights that create artefacts.

We used Metashape 1.8.1 (Agisoft, Inc.) to generate three-dimensional models of panels A and B based on photographs taken with the parameters reported above. Generation of cross-sections and measurements from these models were performed with MeshLab 2021.20. Resolution of the three-dimensional surface is estimated to be accurate to 0.2 mm.